#### Homework-6
##### Total number of points: 40
#### Due date: March 31st, 2023

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on basics of Python. The Python notebooks shared will be helpful to solve these problems. 

Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.

#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

```python
import pandas as pd
import numpy as np
df = pd.read_csv('https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/KaggleAdmissions.csv')
df.columns=['Serial No.', 'GRE Score', 'TOEFL Score', 'University Rating', 'SOP',
       'LOR', 'CGPA', 'Research', 'Chance of Admit']
df.head()
```

Output:
```
Serial No.  GRE Score  TOEFL Score  University Rating  SOP  LOR  CGPA  \
0           1        337          118                  4  4.5  4.5  9.65   
1           2        324          107                  4  4.0  4.5  8.87   
2           3        316          104                  3  3.0  3.5  8.00   
3           4        322          110                  3  3.5  2.5  8.67   
4           5        314          103                  2  2.0  3.0  8.21   

   Research  Chance of Admit  
0         1             0.92  
1         1             0.76  
2         1             0.72  
3         1             0.80  
4         0             0.65
```

### Q1 [18 points]

#### 1. Create a new dataframe `df_sub`, that is a copy of `df` and drop the column `Serial No.` 

#### 2. Standardize only these attributes in `df_sub` using the function `RobustScaler()`: 
- `GRE Score`, `TOEFL Score`

#### 3. Perform normalization only on these attributes in `df_sub` using the function `StandardScaler()`
- `University Rating`, `SOP`, `LOR`, `CGPA`, `Research` 
    
#### Note that after steps 3 and 4, make sure you still have the transformed values saved in `df_sub` 

#### 4. Create a new column named `Admit` using original column `Chance of Admit ` to create a discrete set of class labels using these conditions. Then drop `Chance of Admit ` column from `df_sub`. 
- Convert to 2, if `Chance of Admit` value is `>= 0.65`
- 1, if `Chance of Admit` value is `< 0.65`



<!-- grade_id: Q1-Sol, points: 0, grade: false, solution: true -->
```python
#Answering the above questions in the same order as listed

### BEGIN SOLUTION
from sklearn import preprocessing 


df_sub = df.drop(columns=['Serial No.'])
colnames=list(df_sub.columns)

df_sub[['GRE Score', 'TOEFL Score']]=preprocessing.RobustScaler().fit_transform(df_sub[['GRE Score', 'TOEFL Score']])
df_sub[[ 'University Rating', 'SOP', 'LOR', 'CGPA', 'Research']]=preprocessing.StandardScaler().fit_transform(df_sub[[ 'University Rating', 'SOP', 'LOR', 'CGPA', 'Research']])

bins=np.array([0,0.65,1.0])
df_sub['Admit'] = np.digitize(df['Chance of Admit'], bins)
df_sub = df_sub.drop(columns=['Chance of Admit'])

#print(len(df_sub))
#print(len(df_sub.columns))
#print(set(df_sub.columns))

#print(round(np.mean(df_sub['CGPA']), 2))
#print(round(np.std(df_sub['CGPA']), 2))
#print(round(np.mean(df_sub['TOEFL Score']), 2))


#print(round(np.sum(df_sub['Admit']), 2))
#print(set(df_sub['Admit']))
#print(len(df_sub['Admit'].loc[df_sub['Admit']==1]))
### END SOLUTION

```

<!-- grade_id: 1, points: 6, grade: true, solution: false -->
```python
#[6 points] Test cell-1
#DO NOT MODIFY/DELETE THIS CELL 
assert (len(df_sub))==400
assert (len(df_sub.columns))==8
assert (set(df_sub.columns))=={'LOR', 'SOP', 'CGPA', 'TOEFL Score', 'GRE Score', 'Research', 'Admit', 'University Rating'}

```

<!-- grade_id: 2, points: 6, grade: true, solution: false -->
```python
#[6 points] Test cell-2
#DO NOT MODIFY/DELETE THIS CELL 
assert (round(np.mean(df_sub['CGPA']), 2))==0.0
assert (round(np.std(df_sub['CGPA']), 2))==1.0
assert (round(np.mean(df_sub['TOEFL Score']), 2))==0.05
```

<!-- grade_id: 3, points: 6, grade: true, solution: false -->
```python
#[6 points] Test cell-3
#DO NOT MODIFY/DELETE THIS CELL 
assert (round(np.sum(df_sub['Admit']), 2))==687
assert (set(df_sub['Admit']))=={1, 2}
assert (len(df_sub['Admit'].loc[df_sub['Admit']==1]))==113
```

### Q2 [15 points]
#### 1. Split the data saved in the dataframe in Question-1 (`df_sub`) into `X` and `y` for feature columns and class label column respectively. 
- Feature columns (X): `GRE Score`, `TOEFL Score`, `University Rating`, `SOP`, `LOR`, `CGPA`, `Research` 
- CLass label column (y): `Admit`

#### 2. Using `X` and `y` variables representing features and class labels, perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`). 
- Use test_size=0.4, random_state=55 as the parameters for train_test_split() function. 

#### 3. Train the randomforest classifier (initialized as variable `rf`) using these parameters: max_depth=7, random_state=23. 
- Using the trained model `rf` to first compute accuracy score and assign it to variable `acc_rf`. 
- Then compute the impurity-based feature importances.
- Append the names of these top-3 features to a list `impFeatures`. Please make sure you type the feature names exactly as in `df_sub` 

#### 4. Train a K-Nearest Neighbors classifier (initialized as variable `knn1` and `knn2`) using these parameters: n_neighbors=5 and n_neighbors=20 and the `kd_tree` algorithm.
- Using the trained models `knn1` and `knn2` on training data, compute the accuracy score using test data and assign that to variables `acc_knn_5` and `acc_knn_22` respectively for `k=2` and `k=22`. 

<!-- grade_id: Q2-Sol, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
from sklearn.ensemble import RandomForestClassifier
from sklearn.neighbors import KNeighborsClassifier
from sklearn.metrics import accuracy_score

#Answering the above questions in the same order as listed

impFeatures=[]

### BEGIN SOLUTION
X = df_sub.drop(columns=['Admit'])
y = df_sub['Admit']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.4, random_state=55)

rf = RandomForestClassifier(max_depth=7, random_state=23)
rf.fit(X_train, y_train)
y_pred_rf = rf.predict(X_test)
acc_rf = accuracy_score(y_test, y_pred_rf)
#print(rf.max_depth)
#print(rf.n_estimators)
#print(round(acc_rf,2))

colnames = list(X.columns)
fts  = dict(zip(colnames, rf.feature_importances_))
import operator
sorted_op = sorted(fts.items(), key=operator.itemgetter(1), reverse=True)
for i in range(3):
    impFeatures.append(sorted_op[i][0])

#print(set(impFeatures))

knn1 = KNeighborsClassifier(n_neighbors=5, algorithm='kd_tree')
knn1.fit(X_train, y_train)
y_pred_knn1 = knn1.predict(X_test)
acc_knn_5 = accuracy_score(y_test, y_pred_knn1)
#print(knn1.metric)
#print(round(acc_knn_4,2))


knn2 = KNeighborsClassifier(n_neighbors=22, algorithm='kd_tree')
knn2.fit(X_train, y_train)
y_pred_knn2 = knn2.predict(X_test)
acc_knn_22 = accuracy_score(y_test, y_pred_knn2)
#print(knn1.metric)
#print(round(acc_knn_20, 2))

### END SOLUTION 
```

<!-- grade_id: 4, points: 8, grade: true, solution: false -->
```python
#[8 points] Test cell-4
#DO NOT MODIFY/DELETE THIS CELL 
assert (rf.max_depth)==7
assert (rf.n_estimators)==100
#assert round(acc_rf,2)==0.9
assert round(acc_rf,2)<=0.9
assert set(impFeatures)=={'GRE Score', 'CGPA', 'TOEFL Score'}
```

<!-- grade_id: 5, points: 6, grade: true, solution: false -->
```python
#[6 points] Test cell-5
#DO NOT MODIFY/DELETE THIS CELL 
assert knn1.algorithm=="kd_tree"
assert knn1.algorithm=="kd_tree"
assert (round(acc_knn_5,2))<=0.87
assert (round(acc_knn_22, 2))<=0.9

```

### Q3 [10 points]

#### 1. Assign your response to this string variable `your_response1` explaining why there is a performance difference between `knn1` and `knn2` models that were trained with neighbors `4` and `20` respectively? Justify. 
#### 2. Include your response in this variable `your_response2` describing if this is a reasonable way to perform normalization? 

<!-- grade_id: Q3-Sol, points: 0, grade: false, solution: true -->
```python
your_response1=" "
your_response2=" "

### BEGIN SOLUTION

your_response1=" Larger the value of k, better the accuracy "

your_response2 = "It is interesting but somehow the model that standardizes the entire data might perform better than this "

### END SOLUTION 
```

<!-- grade_id: 6, points: 8, grade: true, solution: false -->
```python
#[8 points] Hidden Test cell-6
#DO NOT MODIFY/DELETE THIS CELL 
### BEGIN HIDDEN TESTS
assert len(your_response1)>4
assert len(your_response2)>4
### END HIDDEN TESTS
```