##### Homework-4
##### Total number of points: 50
#### Due date: February 24th, 2023

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on data manipulation, manipulating strings and feature preprocessing. The Python notebooks shared will be helpful to solve these problems.  

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### Q1 [10 points]. Given the groundtruth or actual class labels `y_true` and the predicted class labels `y_pred`, 
#### build a confusion matrix without using any existing library to compute 
- `fp` false positives, 
- `tp` true positives, 
- `fn` false negatives and 
- `tn` true negatives. 


<!-- grade_id: Q1-sol, points: 0, grade: false, solution: true -->
```python

y_true=[1, 1, 0, 0, 1, 1, 1, 1]
y_pred=[1, 0, 0, 1, 1, 0, 1, 1]

tp=0
fp=0
fn=0
tn=0

### BEGIN SOLUTION
for i in range(len(y_true)):
    if y_true[i]==1 and y_pred[i]==1: 
        tp=tp+1
    elif y_true[i]==0 and y_pred[i]==1:
        fp=fp+1
    elif y_true[i]==1 and y_pred[i]==0:
        fn=fn+1
    elif y_true[i]==0 and y_pred[i]==0:
        tn=tn+1
### END SOLUTION

```

<!-- grade_id: Q1-tests, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert tp==4
assert fp==1
```

<!-- grade_id: Q1-tests2, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
### BEGIN HIDDEN TESTS
assert tn==1
assert fn==2
assert ((tp+tn)/(tp+tn+fp+fn))==0.625
### END HIDDEN TESTS

```

#### Q2 [5 points]. Given the `y_true` which is the groundtruth (or the actual class labels), create a list of predicted labels `y_pred` 
#### that reflects the respective worst-performing model where accuracy will be '0'. 



<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import accuracy_score

y_true = [1, 1, 1, 1, 1, 1, 1, 1]

### BEGIN SOLUTION
y_pred = [0, 0, 0, 0, 0, 0, 0, 0]
### END SOLUTION
```

<!-- grade_id: Q2-tests, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert accuracy_score(y_true, y_pred)==0
```

#### Q3 [10 points]. More Calculations of Classification Performance 

By utilizing the proposed `y_prob` and `y_true` and using a threshold of 0.5 (where the class is 1 if the probability is >= 0.5 and 0 if the probability is < 0.5) calculate the predicted class `y_pred`. Compute the accuracy as `acc`, true positive rate as `tpr`, sensitivity value as `sens`, and Area Under the Receiver Operating Characteristic Curve (ROC AUC) as `auc`. 

**Hint, one of the evaluation functions takes the probability as an input.**


<!-- grade_id: Q3-sol, points: 0, grade: false, solution: true -->
```python
y_true=[0, 1, 0, 0, 1, 1, 1, 0]
y_prob=[0.40, 0.70, 0.20, 0.30, 0.40, 0.65, 0.70, 0.80]


### BEGIN SOLUTION
from sklearn.metrics import confusion_matrix
from sklearn.metrics import roc_auc_score
y_pred=[]
for x in y_prob:
    if x>=0.5:
        y_pred.append(1)
    else:
        y_pred.append(0)
print(y_pred)

CM = confusion_matrix(y_true, y_pred)

tn = CM[0][0]
fn = CM[1][0]
tp = CM[1][1]
fp = CM[0][1]

tpr=(tp)/(tp+fn)
sens=(tp)/(tp + fn)
acc=accuracy_score(y_true, y_pred)
auc=roc_auc_score(y_true, y_prob)
print(acc, round(tpr,3), round(sens,3), round(auc,3))

### END SOLUTION
```

Output:
```
[0, 1, 0, 0, 0, 1, 1, 1]
0.75 0.75 0.75 0.719
```

<!-- grade_id: Q3-tests1, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert y_pred == [0, 1, 0, 0, 0, 1, 1, 1]
assert round(acc,3)==0.75
assert round(sens,3)==0.75
assert round(tpr,3)==0.75
```

<!-- grade_id: Q3-tests2, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert round(auc,3)== 0.719
```

```python
#Please do not modify/delete this cell
import pandas as pd
df=pd.read_csv('https://raw.githubusercontent.com/rpi-techfundamentals/website_fall_2021/master/site/public/titanic_processed_train3.csv')
#View the dataframe
df
```

Output:
```
Age  SibSp  Parch     Fare  Pclass_2  Pclass_3  Sex_male  \
0    22.000000    1.0    0.0   7.2500         0         1         1   
1    38.000000    1.0    0.0  71.2833         0         0         0   
2    26.000000    0.0    0.0   7.9250         0         1         0   
3    35.000000    1.0    0.0  53.1000         0         0         0   
4    35.000000    0.0    0.0   8.0500         0         1         1   
..         ...    ...    ...      ...       ...       ...       ...   
886  27.000000    0.0    0.0  13.0000         1         0         1   
887  19.000000    0.0    0.0  30.0000         0         0         0   
888  29.699118    1.0    2.0  23.4500         0         1         0   
889  26.000000    0.0    0.0  30.0000         0         0         1   
890  32.000000    0.0    0.0   7.7500         0         1         1   

     Embarked_Q  Embarked_S  Survived  pred_prob  pred  
0             0           1         0        0.1     0  
1             0           0         1       
```

### Q4 [10 points]. Evaluation

Calculate the accuracy (`tacc`), precision(`tprec`), recall(`trecall`), and Area Under the Receiver Operating Characteristic Curve (ROC AUC) `tauc` for the dataframe `df`.  The predicted class is in the column `pred` and the predicted probability from the model is in the column `pred_prob`. 

<!-- grade_id: q4-sol, points: 0, grade: false, solution: true -->
```python
#Answer here



### BEGIN SOLUTION
from sklearn.metrics import precision_score
from sklearn.metrics import recall_score
tacc=accuracy_score(df['Survived'], df['pred'])
tprec=precision_score(df['Survived'], df['pred'])
trecall=recall_score(df['Survived'], df['pred'])
tauc=roc_auc_score(df['Survived'], df['pred_prob'])
print(tacc,tprec,trecall,tauc)
### END SOLUTION

```

Output:
```
0.7351290684624018 0.7345132743362832 0.4853801169590643 0.8007941073083438
```

<!-- grade_id: q4-test1, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert tacc>0 and tacc<1.0
assert tprec>0 and tprec<1.0
assert trecall>0 and trecall<1.0
### BEGIN HIDDEN TESTS
assert round(tacc,4)==0.7351
assert round(tprec,4)==0.7345
assert round(trecall,4)==0.4854
### END HIDDEN TESTS


```

<!-- grade_id: q4-test2, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
### BEGIN HIDDEN TESTS
assert round(tauc,4)==0.8008
### END HIDDEN TESTS
```

### Q5 Generate a Confusion Matrix

Use the the standard module confusion_matrix function from Scikit Learn to find the confusion matrix, setting it equal to `cm`. 

Parse the confusion matrix to set the true postives equal to `cm_tp` and the true negatives equal to `cm_tn`. 

Documentation 
[https://scikit-learn.org/0.16/modules/generated/sklearn.metrics.confusion_matrix.html](https://scikit-learn.org/0.16/modules/generated/sklearn.metrics.confusion_matrix.html)




<!-- grade_id: q5-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import confusion_matrix

### BEGIN SOLUTION
cm=confusion_matrix(df['Survived'], df['pred'])
cm_tn=cm[0][0]
cm_tp=cm[1][1]
print(tp,tn, '\n',cm)
print(len(cm))
### END SOLUTION

```

Output:
```
3 3 
 [[489  60]
 [176 166]]
2
```

<!-- grade_id: q5-test, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert len(cm)==2
assert cm_tp>0
assert cm_tn>0
### BEGIN HIDDEN TESTS
assert cm[0][0]==489
assert cm[1][1]==166
assert cm_tn==489
assert cm_tp==166
### END HIDDEN TESTS


```

### q6 Evaluation - By Gender

Conduct the same analysis as above, but for the different subgroups of men and women.

##### Calculate Metrics for Men
Calculate accuracy (`tacc_m`) and Area Under the Receiver Operating Characteristic Curve (ROC AUC) `tauc_m` for only the men.


##### Calculate Metrics for Women
Calculate accuracy (`tacc_w`) and Area Under the Receiver Operating Characteristic Curve (ROC AUC) `tauc_w` for only the women.

<!-- grade_id: q6-sol, points: 0, grade: false, solution: true -->
```python
#Your solution here. 

### BEGIN SOLUTION
df_m=df[df['Sex_male']==1]
df_fm=df[df['Sex_male']==0]
tacc_m=accuracy_score(df_m['Survived'], df_m['pred'])
tacc_fm=accuracy_score(df_fm['Survived'], df_fm['pred'])
tauc_m=roc_auc_score(df_m['Survived'], df_m['pred_prob'])
tauc_fm=roc_auc_score(df_fm['Survived'], df_fm['pred_prob'])
print(tacc_m,tacc_fm,tauc_m,tauc_fm)
### END SOLUTION
```

Output:
```
0.7989601386481803 0.6178343949044586 0.7364051595702972 0.7651141842844275
```

<!-- grade_id: q6-tests1, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
assert len(cm)==2
assert tacc_m>0
assert tacc_fm>0
### BEGIN HIDDEN TESTS
assert round(tacc_m,2)==0.80
assert round(tacc_fm,2)==0.62
### END HIDDEN TESTS
```

<!-- grade_id: q6-tests2, points: 5, grade: true, solution: false -->
```python
#Please do not modify/delete this cell
### BEGIN HIDDEN TESTS
assert round(tauc_m,2)==0.74
assert round(tauc_fm,2)==0.77
### END HIDDEN TESTS
```