###### Exam 1: Spring 2023
#### Total number of points: 40

This is an open-book test. However, talking or chatting with your classmates is NOT permitted and will be considered plagiarism. 

<b>Before you submit this test, make sure your answers are saved, and everything runs as expected</b>. 

- For empty notebook submissions, you will receive 0 points. 
- If instructions (including not using the same variable names) are not followed that led to a wrong answer, you will receive 0 points. 

First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).
Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your notebook at the appropriate link to submit this homework. Make sure your answers are saved before you submit. 

#### Please note that for any question in this test you will receive points ONLY if your solution passes all the hidden testcases. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### For this test, we will be using Stroke data provided by WHO. The dataset is used to predict whether a patient is likely to get stroke based on the input parameters like gender, age, various diseases, and smoking status 
<ul> 
    <li>id: unique identifier</li>
    <li>gender: "Male", "Female" or "Other"</li>
    <li>age: age of the patient</li>
    <li>hypertension: 0 if the patient doesn't have hypertension, 1 if the patient has hypertension</li>
    <li>heart_disease: 0 if the patient doesn't have any heart diseases, 1 if the patient has a heart disease</li>
    <li>ever_married: "No" or "Yes"</li>
    <li>work_type: 'Private', 'Self-employed', 'Govt_job', 'children', 'Never_worked'</li>
    <li>Residence_type: "Rural" or "Urban"</li>
    <li>avg_glucose_level: average glucose level in blood</li>
    <li>bmi: body mass index</li>
    <li>smoking_status: "formerly smoked", "never smoked", "smokes" or "Unknown"*</li>
    <li>stroke: 1 if the patient had a stroke or 0 if not -- <b>Class label column</b> </li>
</ul>
*Note: "Unknown" in smoking_status means that the information is unavailable for this patient

```python
#Please run this cell before answering the following questions 
import numpy as np
import pandas as pd

df = pd.read_csv('https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/KaggleStroke.csv')
df.head()
```

Output:
```
id  gender   age  hypertension  heart_disease ever_married  \
0   9046    Male  67.0             0              1          Yes   
1  51676  Female  61.0             0              0          Yes   
2  31112    Male  80.0             0              1          Yes   
3  60182  Female  49.0             0              0          Yes   
4   1665  Female  79.0             1              0          Yes   

       work_type Residence_type  avg_glucose_level   bmi   smoking_status  \
0        Private          Urban             228.69  36.6  formerly smoked   
1  Self-employed          Rural             202.21   NaN     never smoked   
2        Private          Rural             105.92  32.5     never smoked   
3        Private          Urban             171.23  34.4           smokes   
4  Self-employed          Rural             174.12  24.0     never smoked   

   stroke  
0       1  
1       1  
2       1  
3       1  
4       1
```

### Q-1 [10 points] Please answer the following questions in the same order: 

1. Assign variable `rows` and `cols` with the total number of data points and the number of attributes present in the DataFrame df.
2. Assign the list `colnames` with the names of features present in this DataFrame df.
3. Find the value of `age` associated with the sample whose `id` is `1665` and assign this value to variable `sid1665_age`.
4. Find the total number of missing values in 'smoking_status' column and assign it to variable `sstat_mvals`.
5. Remove the column `id` from the DataFrame `df`. Make sure this is reflected in the DataFrame `df`.

<!-- grade_id: Q1-Sol, points: 0, grade: false, solution: true -->
```python
#Please make sure you use the exact names for variables including the casing of letters

### BEGIN SOLUTION
rows=df.shape[0]
cols=df.shape[1]
colnames = list(df.columns)
sid1665_age = df.loc[df['id']==1665,'age'].values[0]
sstat_mvals = df['smoking_status'].isnull().values.sum()
df = df.drop(columns=['id'])

#print(rows)
#print(cols)
#print(set(colnames))
#print(sid1665_age)
#print(sstat_mvals)
#print(len(df.columns))
### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 
assert 'rows' in globals()
assert 'cols' in globals()
assert 'colnames' in globals()
assert 'sid1665_age' in globals()
assert 'sstat_mvals' in globals()

```

<!-- grade_id: 1, points: 1, grade: true, solution: false -->
```python
#Test cell-1 [1pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert cols==12
### END HIDDEN TESTS

```

<!-- grade_id: 2, points: 1, grade: true, solution: false -->
```python
#Test cell-2 [1pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert rows==5110
### END HIDDEN TESTS

```

<!-- grade_id: 3, points: 2, grade: true, solution: false -->
```python
#Test cell-3 [2pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert set(colnames)=={'work_type', 'ever_married', 'avg_glucose_level', 'smoking_status', 'stroke', 'gender', 'hypertension', 'bmi', 'Residence_type', 'id', 'heart_disease', 'age'}
### END HIDDEN TESTS

```

<!-- grade_id: 4, points: 2, grade: true, solution: false -->
```python
#Test cell-4 [2pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert sid1665_age==79.0
### END HIDDEN TESTS

```

<!-- grade_id: 5, points: 2, grade: true, solution: false -->
```python
#Test cell-5 [2pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert sstat_mvals==0
### END HIDDEN TESTS


```

<!-- grade_id: 6, points: 2, grade: true, solution: false -->
```python
#Test cell-6 [2pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert len(df.columns)==11
### END HIDDEN TESTS
```

### Q-2 [15 points] Please answer the following questions in the same order: 

1. Perform one-hot encoding on the DataFrame `df` using the Pandas `get_dummies()` function with default parameters. Assign this newly transformed DataFrame as `newdf`.
2. Replace all the missing values in this transformed DataFrame `newdf` with the average value of the corresponding column.
3. Split the DataFrame `newdf` into `X` that includes all the feature columns, and `y` that includes ONLY the class label column which is `stroke`. Please follow the case format of variables used here -- otherwise, your code wont pass the testcases.

More details about the `get_dummies()` function can be found here: 
<a href="https://pandas.pydata.org/docs/reference/api/pandas.get_dummies.html"> Link </a>



<!-- grade_id: Q2-Sol, points: 0, grade: false, solution: true -->
```python
#Please make sure you use the exact names for variables including the casing of letters

### BEGIN SOLUTION
newdf = pd.get_dummies(df)
#print(len(df.columns))
#print(len(newdf.columns))
#print(newdf.isnull().values.sum())

#Replacing all the missing values with the average value of the corresponding column
mean1 = newdf['bmi'].mean()
newdf['bmi'] = newdf['bmi'].fillna(mean1)
#print(round(newdf['gender_Female'].mean(), 2))
#print(newdf.isnull().values.any())
#print(round(newdf['bmi'].mean(), 2))

#Splitting the data into X and y 
X = newdf.drop(columns=['stroke'])
y = newdf['stroke']
#print(newdf.isnull().values.sum())
#print(len(X.columns))
#print(round(y.mean(), 2))
### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 
assert 'newdf' in globals()
assert 'X' in globals()
assert 'y' in globals()
```

<!-- grade_id: 7, points: 3, grade: true, solution: false -->
```python
#Test cell-7 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert len(df.columns)==11
assert len(newdf.columns)==22
### END HIDDEN TESTS
```

<!-- grade_id: 8, points: 3, grade: true, solution: false -->
```python
#Test cell-8 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert (newdf.isnull().values.sum())==0
### END HIDDEN TESTS
```

<!-- grade_id: 9, points: 3, grade: true, solution: false -->
```python
#Test cell-9 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert round(newdf['gender_Male'].mean(), 2)==0.41
assert round(newdf['gender_Female'].mean(), 2)==0.59
### END HIDDEN TESTS
```

<!-- grade_id: 10, points: 3, grade: true, solution: false -->
```python
#Test cell-10 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert newdf.isnull().values.any()==False
assert round(newdf['bmi'].mean(), 2)==28.89
### END HIDDEN TESTS
```

<!-- grade_id: 11, points: 3, grade: true, solution: false -->
```python
#Test cell-11 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert len(X.columns)==21
assert round(y.mean(), 2)==0.05
### END HIDDEN TESTS
```

### Q-3 [15 points] Please answer the following questions in the same order: 

1. Perform normalization of the features `X`, using `MinMaxScaler()` with default parameters, and assign the result to the variable `X_std`.
2. Create a list of values `y_pred` of length `len(y)` with all values set to 1.
3. Compute accuracy using the ground truth labels `y` and fake predicted labels `y_pred` and save it as a variable `accu`.
4. Compute the precision and recall values using `y` and `y_pred` and assign these values to the variables `prec` and `recall`, respectively.
5. Compute the F1 score and assign it to variable `fscore`.




<!-- grade_id: Q3-Sol, points: 0, grade: false, solution: true -->
```python
#Please make sure you use the exact names for variables including the casing of letters

### BEGIN SOLUTION

#Normalization 
from sklearn.preprocessing import MinMaxScaler

X_std = MinMaxScaler().fit_transform(X)
#print(X_std)
#print(len(X_std))
#print(len(X_std[0]))

#Generate a random list with 0s and 1s 
y_pred = [1 for i in range(len(y))]
#print(len(y_pred))
#print(set(y_pred))

#Accuracy 
from sklearn.metrics import accuracy_score
accu = accuracy_score(y, y_pred)
#print(round(accu, 2))

#Precision and Recall
from sklearn.metrics import precision_score
from sklearn.metrics import recall_score

prec = precision_score(y, y_pred)
#print(round(prec, 2))

recall = recall_score(y, y_pred)
#print(round(recall, 2))

#Fscore
fscore = (2.0*prec*recall)/(prec+recall)
#print(round(fscore, 2))


### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 

assert 'y_pred' in globals()
assert 'accu' in globals()
assert 'prec' in globals()
assert 'recall' in globals()
assert 'fscore' in globals()
assert 'X_std' in globals()
```

<!-- grade_id: 12, points: 3, grade: true, solution: false -->
```python
#Test cell-12 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert len(X_std)==5110
assert len(X_std[0])==21
### END HIDDEN TESTS
```

<!-- grade_id: 13, points: 3, grade: true, solution: false -->
```python
#Test cell-13 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert len(y_pred)==5110
assert set(y_pred)=={1}
### END HIDDEN TESTS
```

<!-- grade_id: 14, points: 3, grade: true, solution: false -->
```python
#Test cell-14 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert round(accu, 2)==0.05
### END HIDDEN TESTS
```

<!-- grade_id: 15, points: 3, grade: true, solution: false -->
```python
#Test cell-15 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert round(prec, 2)==0.05
assert round(recall, 2)==1.0
### END HIDDEN TESTS
```

<!-- grade_id: 16, points: 3, grade: true, solution: false -->
```python
#Test cell-16 [3pts] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 


### BEGIN HIDDEN TESTS
assert round(fscore, 2)==0.09
### END HIDDEN TESTS
```