#### Homework-3
##### Total number of points: 60
#### Due date: February 14, 2023

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on data manipulation, manipulating strings and feature preprocessing. The Python notebooks shared will be helpful to solve these problems.  

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

Use the titanic dataset from this url ('https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/titanictrain.csv') and we will use the same dataset in some of the questions here below. 

#### Q1. Shape of a Data Frame. 
1. Import pandas package as pd.
2. Read the file from the `url` and load the data into dataframe `df` with default index.
3. Set number of rows equal to the variable `rows` and the number of columns equal to the variable `cols`. 
4. Print out the number of rows and columns, clearly labeling each.



<!-- grade_id: Q1-Sol, points: 0, grade: false, solution: true -->
```python

url='https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/titanictrain.csv'

### BEGIN SOLUTION
import pandas as pd
df = pd.read_csv(url)
rows=df.shape[0]
cols=df.shape[1]
cols1=len(df.iloc[1,:])
### END SOLUTION
```

<!-- grade_id: 1, points: 3, grade: true, solution: false -->
```python
#[C1] DO NOT MODIFY/DELETE THIS CELL
assert rows==891
assert df['Age'].sum()==21205.17

```

<!-- grade_id: 2, points: 3, grade: true, solution: false -->
```python
#[C2]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert cols==12
assert df.isnull().sum().sum()==866
### END HIDDEN TESTS
```

Now we will be using the above dataframe df to do some preprocessing operations on this data. All the required libraries for further processing will be loaded here. 

```python
import numpy as np
import pandas as pd
import re
from sklearn.model_selection import train_test_split
```

#### Q2. Dataframe Basic Analyses
Determine how many first, second, and third class (assiging to the variables `class1`, `class2`, `class3`) passangers there are. 

*Hint -- Use value_counts operation* 



<!-- grade_id: Q2-sol, points: 0, grade: false, solution: true -->
```python
#Your answer here. 

### BEGIN SOLUTION
class1 = df['Pclass'].value_counts()[1]
class2 = df['Pclass'].value_counts()[2]
class3 = df['Pclass'].value_counts()[3]
#print(ggby)
### END SOLUTION

```

<!-- grade_id: 3, points: 3, grade: true, solution: false -->
```python
#[C3] DO NOT MODIFY/DELETE THIS CELL
assert class1>0
assert class2>0
assert class3>0


```

<!-- grade_id: 4, points: 3, grade: true, solution: false -->
```python
#[C4] DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert class1==216
assert class2==184
assert class3==491
### END HIDDEN TESTS
```

#### Q3. Groupby

Now use a groupby statement to calculate the mean age (use the 'age' attribute) of passengers who are of different gender. 

Round the age to 2 decimal places (for example 3.14156 converts to 3.14 ) and assign the resulting variable to `female_age`, `male_age`.



<!-- grade_id: Q3-Sol, points: 0, grade: false, solution: true -->
```python
#Your answer 

### BEGIN SOLUTION
female_age = round(df.groupby(['Sex']).Age.mean()['female'], 2)
male_age = round(df.groupby(['Sex']).Age.mean()['male'], 2)


### END SOLUTION
```

<!-- grade_id: 5, points: 3, grade: true, solution: false -->
```python
# [C5]DO NOT MODIFY/DELETE THIS CELL
assert female_age>0
assert isinstance(female_age, np.floating)
assert male_age>0 
assert isinstance(male_age, np.floating)


```

<!-- grade_id: 6, points: 3, grade: true, solution: false -->
```python
# [C6]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert female_age==27.92
assert male_age==30.73
### END HIDDEN TESTS
```

#### Q4. Split Dataframe

Now split the dataframe `df` into 3 different dataframes `dfclass1`, `dfclass2`, and `dfclass3` using the `Pclass` variable. 



<!-- grade_id: Q4-sol, points: 0, grade: false, solution: true -->
```python
#Your answer

### BEGIN SOLUTION
dfclass1 = df.loc[df['Pclass']==1,:]
dfclass2 = df.loc[df['Pclass']==2,:]
dfclass3 = df.loc[df['Pclass']==3,:]

### END SOLUTION

```

<!-- grade_id: 7, points: 2, grade: true, solution: false -->
```python
#[C7]DO NOT MODIFY/DELETE THIS CELL
assert len(dfclass1)>0
assert len(dfclass2)>0
assert len(dfclass3)>0
assert isinstance(dfclass1, pd.DataFrame)
assert isinstance(dfclass2, pd.DataFrame)
assert isinstance(dfclass3, pd.DataFrame)

```

<!-- grade_id: 8, points: 2, grade: true, solution: false -->
```python
#[C8]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfclass1.shape[0]==216
### END HIDDEN TESTS
```

<!-- grade_id: 9, points: 2, grade: true, solution: false -->
```python
#[C9]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfclass2.shape[0]==184
### END HIDDEN TESTS
```

<!-- grade_id: 10, points: 2, grade: true, solution: false -->
```python
#[C10]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfclass3.shape[0]==491
### END HIDDEN TESTS
```

#### Q5. Filter Missing Values

Create a new dataframe `dfna1` from `df` which removes all rows in which *any* of the variables missing.

Create a new dataframe `dfna2` from `df` which removes all rows in which *all* of the variables columns `Age` and `Cabin`.

<!-- grade_id: Q5-Sol, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
dfna1 = df.dropna()
dfna2 = df.dropna(how='all',subset=['Age','Cabin'])
### END SOLUTION

```

<!-- grade_id: 11, points: 2, grade: true, solution: false -->
```python
#[C11]DO NOT MODIFY/DELETE THIS CELL
assert len(dfna1)>0
assert len(dfna2)>0
assert isinstance(dfna1, pd.DataFrame)
assert isinstance(dfna2, pd.DataFrame)
```

<!-- grade_id: 12, points: 2, grade: true, solution: false -->
```python
#[C12]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfna1.shape[0]==183
### END HIDDEN TESTS
```

<!-- grade_id: 13, points: 2, grade: true, solution: false -->
```python
#[C13]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfna2.shape[0]==733
### END HIDDEN TESTS
```

#### Q6. Stratified sampling 

Utilize the original dataframe 'df'.  

1. Set the `y` variable to the `Survived` column.
2. Set X to the `SibSp`,`Parch`, and `Fare` columns of the dataframe. 
3. Create `train_X`, `test_X`, `train_y`, `test_y` by doing a 50% 50% split with stratification by the `Survived` varaible. Use `random_state` equal to 123. 


<!-- grade_id: Q6-sol, points: 0, grade: false, solution: true -->
```python
#Your answer

### BEGIN SOLUTION
y=df['Survived']
X=df[['SibSp','Parch','Fare']]
train_X, test_X, train_y, test_y = train_test_split(X, y, train_size=0.5,test_size=0.5,random_state=123,stratify=y)
### END SOLUTION
```

<!-- grade_id: 14, points: 2, grade: true, solution: false -->
```python
#[C14]DO NOT MODIFY/DELETE THIS CELL
assert len(train_X)>0
assert len(test_X)>0
assert isinstance(train_X, pd.DataFrame)
assert isinstance(test_X, pd.DataFrame)
assert isinstance(train_y, pd.Series)
assert isinstance(test_y, pd.Series)
```

<!-- grade_id: 15, points: 2, grade: true, solution: false -->
```python
#[C15]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert train_X.shape[0]==445
### END HIDDEN TESTS
```

<!-- grade_id: 16, points: 2, grade: true, solution: false -->
```python
#[C16]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert train_y.sum()==171
### END HIDDEN TESTS
```

<!-- grade_id: 17, points: 2, grade: true, solution: false -->
```python
#[C17]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert test_y.sum()==171
### END HIDDEN TESTS
```

#### Q7.  Feature Manipulation
1. Consider the original dataframe `df`. Create a copy of this and call it `dfage`.
2. Remove all the rows of `dfage` in which the `Age` value is missing. 
3. Then create a new column `dfage['Age_st']` which replaces each value in the 'age' attribute with the corresponding standardized value *rounded to 2 decimal places*. You are not modifying `dfage['Age']`.


*Hint: See this for definition of standardization value in Lecture-8 slides You can use numpy library to compute mean and standard deviation values -- np.mean() and np.std() functions to compute the mean and standard deviation values.*

<!-- grade_id: Q7-sol, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
dfage=df.copy()
dfage = dfage.loc[df['Age'].notna(),]
meanval = np.mean(dfage['Age'])
stdval=np.std(dfage['Age'])
dfage['Age_st'] = dfage['Age'].apply(lambda x:round((x-meanval)/stdval,2))
### END SOLUTION
```

<!-- grade_id: 18, points: 4, grade: true, solution: false -->
```python
#[C18]DO NOT MODIFY/DELETE THIS CELL
assert len(dfage)>0
assert isinstance(dfage, pd.DataFrame)
assert 'Age_st' in dfage.columns


```

<!-- grade_id: 19, points: 2, grade: true, solution: false -->
```python
#[C19]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert dfage['Age_st'][3]==0.37
### END HIDDEN TESTS
```

<!-- grade_id: 20, points: 2, grade: true, solution: false -->
```python
#[C20]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert len(dfage)==714
### END HIDDEN TESTS
```

<!-- grade_id: 21, points: 2, grade: true, solution: false -->
```python
#[C21]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert round(dfage['Age_st'].sum(),2)==0.14
### END HIDDEN TESTS
```

#### Q8. Feature Creation 
1. Create a copy of df called `df8`.
2. In `df8` create a feature called `stown` using `Embarked`.  

    For the `Embarked` is 'S' make `stown` 1. 
    
    Otherwise `stown` is 0.

<!-- grade_id: Q8-sol, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
df8=df.copy()
df8['stown']=0
df8.loc[df['Embarked']=='S','stown']=1
### END SOLUTION
```

<!-- grade_id: 22, points: 3, grade: true, solution: false -->
```python
#[C22]DO NOT MODIFY/DELETE THIS CELL
assert set(df8['stown'])=={0,1}
assert df8.shape==(891, 13)

```

<!-- grade_id: 23, points: 3, grade: true, solution: false -->
```python
#[C23]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert df8['stown'].sum()==644
### END HIDDEN TESTS
```

#### Q9. Attribute Manipulation -- categorical to numerical 
1. Create a copy of df called `df9`.
2. In `df9` create a feature called `gendervar` using feature `Sex`.  

    If the value in `Sex` is `Male` convert it to 0. 
    
    Otherwise `Female` to 1.

<!-- grade_id: Q9-Sol, points: 0, grade: false, solution: true -->
```python
# Attribute manipulation

### BEGIN SOLUTION
df9=df.copy()
df9['gendervar']=df9['Sex'].map({'male':0, 'female':1})
#print(set(df9['gendervar']))
#print((df9['gendervar']).sum())
### END SOLUTION
```

<!-- grade_id: 24, points: 2, grade: true, solution: false -->
```python
##[C24]DO NOT MODIFY/DELETE THIS CELL
assert set(df9['gendervar'])=={0,1}
```

<!-- grade_id: 25, points: 2, grade: true, solution: false -->
```python
#[C25]DO NOT MODIFY/DELETE THIS CELL
### BEGIN HIDDEN TESTS
assert (df9['gendervar']).sum()==314
### END HIDDEN TESTS
```