###### Exam-2: Fall 2022
#### Total number of points: 49

This is an open-book test. However, talking or chatting with your classmates is NOT permitted and will be considered plagiarism. 

<b>Before you submit this test, make sure your answers are saved, and everything runs as expected</b>. 

- For empty notebook submissions, you will receive 0 points. 
- If instructions (including not using the same variable names) are not followed that led to a wrong answer, you will receive 0 points. 

First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).
Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your notebook at the appropriate link to submit this homework. Make sure your answers are saved before you submit. 

#### Please note that for any question in this test you will receive points ONLY if your solution passes all the hidden testcases. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

### Q-1 [10 points] Please answer the following questions in the same order as listed here below: 

1. Assign variable `dataSize` and `features` with the total number of data points present in the dataset and the total number of features present in this dataset.
2. Assign the the names of features present in the iris dataset to a list named `colnames`.
3. Convert the data structure `X` into a dataframe `df`.  **Assign the column names as a approiate for the dataframe**.

<!-- grade_id: Q1-sol, points: 0, grade: false, solution: true -->
```python
from sklearn.datasets import load_iris
import pandas as pd

iris = load_iris()
X = iris.data
y = iris.target

### BEGIN SOLUTION
colnames = iris.feature_names
#print(colnames)

dataSize = len(X)
#print(dataSize)
features = len(colnames)
#print(features)

#convert X from an ndarray to a dataframe with column names
df = pd.DataFrame(X, columns=colnames)
#print(df.head())
#print(y)
#print(type(df))
#print(isinstance(df, pd.DataFrame))
#print(df.loc[1].sum())
### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 
assert 'dataSize' in globals()
assert 'features' in globals()
assert 'colnames' in globals()
assert 'X' in globals()
assert 'y' in globals()
assert 'df' in globals()
```

<!-- grade_id: 1, points: 4, grade: true, solution: false -->
```python
#Test cell-1 [4pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert dataSize==150
assert features==4
### END HIDDEN TESTS
```

<!-- grade_id: 2, points: 2, grade: true, solution: false -->
```python
#Test cell-2 [1pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert colnames==['sepal length (cm)', 'sepal width (cm)', 'petal length (cm)', 'petal width (cm)']
### END HIDDEN TESTS
```

<!-- grade_id: 3, points: 4, grade: true, solution: false -->
```python
#Test cell-3 [4pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
### BEGIN HIDDEN TESTS
assert isinstance(df, pd.DataFrame)==True
assert df.loc[1].sum()==9.5
### END HIDDEN TESTS
```

### Q-2 [10 points] Please answer the following questions in the same order as listed here below: 

Use the dataframe `df` as loaded below following questions. 

This data consists of an dependent variable `Income` (`y`) and a variety of features (`X`) related to purchase history. 

1. Cluster the data in dataframe `X` using KMeans algorithm to create 3 clusters. Use a `random_state`=10. Save the Sum of squared distances of samples to their closest cluster center to a variable `C1_sse`.


2. Save the cluster labels in the variable `C1_label`.  



```python
#This loads our data
import pandas as pd
pd.set_option('display.max_columns', None)
df=pd.read_csv('https://raw.githubusercontent.com/rpi-techfundamentals/website_fall_2021/master/site/public/q2.csv' )
X = df.iloc[:,1:22]
y = df['Income']

```

<!-- grade_id: Q2-Sol, points: 0, grade: false, solution: true -->
```python
#Use the dataframe X for answering this question

### BEGIN SOLUTION

#k-means algorithm
from sklearn.cluster import KMeans
from sklearn.preprocessing import StandardScaler
kmeans = KMeans(n_clusters=3, random_state=10).fit(X)
C1_sse = kmeans.inertia_
C1_label=kmeans.labels_
#print(round(C1_sse,0))
### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 
assert 'C1_sse' in globals()
assert 'C1_label' in globals()

```

<!-- grade_id: 4, points: 5, grade: true, solution: false -->
```python
#Test cell-4 [5pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
### BEGIN HIDDEN TESTS
assert round(C1_sse,0)==112838588.0
### END HIDDEN TESTS
```

<!-- grade_id: 5, points: 5, grade: true, solution: false -->
```python
#Test cell-5 [5pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
from sklearn import metrics
kmeans2 = KMeans(n_clusters=3, random_state=10).fit(X)
C1_label_answer=kmeans2.labels_
assert metrics.adjusted_mutual_info_score(C1_label, C1_label_answer)==1.0
### END HIDDEN TESTS
```

### Q-3 [5 points] Please answer the following questions in the same order as listed here below: 

1. Standardize X using the `StandardScaler()` and assign it to a DataFrame X2 where the column names are the same as they are for the origional X. 



<!-- grade_id: Q3-2-SOL, points: 0, grade: false, solution: true -->
```python
#Use the dataframe df created in Question-1 for answering this question

#normalize the dataset 
X2 = StandardScaler().fit_transform(X)
X2=pd.DataFrame(X2, columns =X.columns)

### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 
assert 'X2' in globals()

```

<!-- grade_id: 6a, points: 3, grade: true, solution: false -->
```python
#Test cell-6a [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert X2.mean().sum()<=0.01

### END HIDDEN TESTS
```

<!-- grade_id: 6b, points: 2, grade: true, solution: false -->
```python
#Test cell-6b [2pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 

### BEGIN HIDDEN TESTS
assert isinstance(X2, pd.DataFrame)==True

### END HIDDEN TESTS
```

## Q-4 [12 points] Please answer the following questions in the same order as listed here below: 

Use the values `X_train`, `X_test`, `y_train`, `y_test` provided below to answer the following questions. 
 
2. Now use DecisionTreeRegressor with `random_state`=0 and `max_depth`=`9` to *train using the training dataset* and predict the `Income` using both the training dataset (creating `y_train_pred`) and testing dataset `y_test_pred`. 

3. Compute the r2 for both the training (`r2_train`) and testing (`r2_test`) . 



```python
from sklearn.model_selection import train_test_split
df=pd.read_csv('https://raw.githubusercontent.com/rpi-techfundamentals/website_fall_2021/master/site/public/q2.csv' )
X = df.iloc[:,1:22]
y = df['Income']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.5, random_state=10)

```

<!-- grade_id: Q4-Sol, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor


### BEGIN SOLUTION

dt = DecisionTreeRegressor(max_depth=9, random_state=0)
dt.fit(X_train, y_train)
y_train_pred = dt.predict(X_train)
y_test_pred = dt.predict(X_test)

r2_train=metrics.r2_score(y_train, y_train_pred)
r2_test=metrics.r2_score(y_test, y_test_pred)

### END SOLUTION
```

```python
#This will test that you have not mislabeled a variable. 

assert 'y_train_pred' in globals()
assert 'y_test_pred' in globals()
assert 'r2_train' in globals()
assert 'r2_test' in globals()
```

<!-- grade_id: 7, points: 3, grade: true, solution: false -->
```python
#Test cell-7 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_train_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(y_train_pred.sum(),0)==58408472.0
### END HIDDEN TESTS
```

<!-- grade_id: 8, points: 3, grade: true, solution: false -->
```python
#Test cell-8 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_test_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(y_test_pred.sum(),0)==58710341.0
### END HIDDEN TESTS
```

<!-- grade_id: 9, points: 3, grade: true, solution: false -->
```python
#Test cell-9 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(r2_train,3))
### BEGIN HIDDEN TESTS
assert round(r2_train,3)==0.969
### END HIDDEN TESTS
```

<!-- grade_id: 10, points: 3, grade: true, solution: false -->
```python
#Test cell-10 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(r2_test,3))
### BEGIN HIDDEN TESTS
assert round(r2_test,3)==0.571
### END HIDDEN TESTS
```

## Q-5 [12 points] Please answer the following questions in the same order as listed here below: 
3. Now train a logistic regression classifier (with `random_state`=0 and `solver`=`newton-cg`) to predict whether someone is a high earner (they have an income one standard deviation above the mean.) Train on the training data below and predict using both the training (`y_train_pred_b`) and test (`y_test_pred_b`) sets. 

Compute the accuracy score  of training (`accuracy_train`) and  test (`accuracy_test`). 


```python
from sklearn.model_selection import train_test_split
df=pd.read_csv('https://raw.githubusercontent.com/rpi-techfundamentals/website_fall_2021/master/site/public/q2.csv' )
X = df.iloc[:,1:22]
#Set threshold of mean plus 1 standard devaiation
high_earner_thres=df['Income'].mean()+df['Income'].std()
df['HighEarner']=0
df.loc[df['Income']>=high_earner_thres,'HighEarner']=1
y=df['HighEarner']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.5, random_state=10)

```

<!-- grade_id: Q5, points: 0, grade: false, solution: true -->
```python
from sklearn.linear_model import LogisticRegression
### BEGIN SOLUTION

dt = LogisticRegression(random_state=0, solver='newton-cg')
dt.fit(X_train, y_train)
y_train_pred_b = dt.predict(X_train)
y_test_pred_b = dt.predict(X_test)

accuracy_train=metrics.accuracy_score(y_train, y_train_pred_b)
accuracy_test=metrics.accuracy_score(y_test,y_test_pred_b)
print(y_train_pred_b.sum(),y_test_pred_b.sum(),round(accuracy_train,2),round(accuracy_test,2 ))
### END SOLUTION
```

Output:
```
121 112 0.93 0.91
```

```python
#This will test that you have not mislabeled a variable. 

assert 'y_train_pred_b' in globals()
assert 'y_test_pred_b' in globals()
assert 'accuracy_train' in globals()
assert 'accuracy_test' in globals()
```

<!-- grade_id: 11, points: 3, grade: true, solution: false -->
```python
#Test cell-11 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_train_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(y_train_pred_b.sum(),0)==121
### END HIDDEN TESTS
```

<!-- grade_id: 12, points: 3, grade: true, solution: false -->
```python
#Test cell-12 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_train_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(y_test_pred_b.sum(),0)==112
### END HIDDEN TESTS
```

<!-- grade_id: 13, points: 3, grade: true, solution: false -->
```python
#Test cell-13 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_train_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(accuracy_train.sum(),2)==0.93
### END HIDDEN TESTS
```

<!-- grade_id: 14, points: 3, grade: true, solution: false -->
```python
#Test cell-14 [3pt] Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
#print(round(y_train_pred.sum(),0))
### BEGIN HIDDEN TESTS
assert round(accuracy_test.sum(),2)==0.91
### END HIDDEN TESTS
```