#### Homework-7


##### Total number of points: 40
#### Due date: April 11th, 2023

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on random forests (including feature importances), and neural networks. The Python notebooks shared will be helpful to solve these problems.  

Steps to evaluate your solutions:

Step-1: Try on Colab or Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the homework's .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as ".ipynb"

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.


#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

```python
#Please run this cell to load the dataset. 
url = 'https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/dataset_Facebook.csv'

import pandas as pd
import math
df  = pd.read_csv(url)
df = df.dropna()
df.columns
```

```python
df
```

#### Part-1 [24 points]: We will leverage randomforest classifier to identify the best features that predict 'Total Interactions'.

#### 1. Create a new dataframe `df_sub2` that contains only these feature columns in `df` : 
- `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`,  `Total Interactions`

#### 2. Transform the categorical attribute `Type` in `df_sub2` to numerical attribute this way: 
- `Link`:1, `Photo`:2, `Status`:3, `Video`:4

#### 3. Perform Standardization (using this formula https://en.wikipedia.org/wiki/Standard_score) only on `Page total likes` column in `df_sub2` 
- Please use `<dataframe>['<column>'].mean()` `<dataframe>['<column>'].std()` if you are using the mean and std values to manipulate the column. 
    
#### 4. Using `df_sub2` perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`).
- Use test_size=0.3, random_state=42 as the parameters for train_test_split() function. 
- Feature columns (X): `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`
- CLass label column (y): `Total Interactions`

#### 5. Train the randomforest classifier (initialized as variable `rf`) using these parameters: max_depth=3, random_state=0. 
- Using the trained model `rf` compute the impurity-based feature importances.
- Append the names of these top-3 features to list `impFeatures`. Please make sure you type the feature names exactly as in df_sub2 




<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python


### BEGIN SOLUTION

from sklearn.ensemble import RandomForestClassifier
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score


df_sub2=df.loc[:, ['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid', 'Total Interactions']]
print(set(df_sub2.columns))
print(len(df_sub2))
print(len(df_sub2.iloc[0,:]))

df_sub2['Type']=df_sub2['Type'].map({'Link':1, 'Photo':2, 'Status':3, 'Video':4})
meanval=df_sub2['Page total likes'].mean()
stdval=df_sub2['Page total likes'].std()
df_sub2['Page total likes']=(df_sub2['Page total likes']-meanval)/stdval



X=df_sub2.drop(columns=['Total Interactions'])
y=df_sub2['Total Interactions']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)
print(len(X_train))
print(len(y_test))

rf = RandomForestClassifier(max_depth=3, random_state=0)
rf.fit(X_train, y_train)

feats=(rf.feature_importances_)
colnames=list(X.columns)

diction = dict(zip(colnames, feats))

import operator
sorted_x = sorted(diction.items(), key=operator.itemgetter(1), reverse=True)
impFeatures=[]
impFeatures.append(sorted_x[0][0])
impFeatures.append(sorted_x[1][0])
impFeatures.append(sorted_x[2][0])
print(impFeatures)

print(rf.n_classes_)
print(rf.n_features_)
### END SOLUTION

```

```python
diction
```

<!-- grade_id: q11, points: 6, grade: true, solution: false -->
```python
##Cell-1 -- Do not modify this cell
assert set(df_sub2.columns)=={'Post Hour', 'Total Interactions', 'Post Weekday', 'Category', 'Paid', 'Post Month', 'Page total likes', 'Type'}
assert len(df_sub2)==495
assert len(df_sub2.iloc[0,:])==8

```

<!-- grade_id: q12, points: 6, grade: true, solution: false -->
```python
##Cell-2 -- Do not modify this cell
assert round(df_sub2['Page total likes'].mean(),0)==0
assert len(X_train)==346
assert 'Post Hour' in impFeatures

```

<!-- grade_id: q13, points: 6, grade: true, solution: false -->
```python
##Cell-3 -- Do not modify this cell
assert rf.n_classes_==230
assert len(impFeatures)==3
assert math.ceil(df_sub2['Page total likes'].std())==1

```

<!-- grade_id: q14, points: 6, grade: true, solution: false -->
```python
##Cell-4 -- Do not modify this cell
assert (rf.n_features_in_)==7
assert len(y_test)==149
assert set(impFeatures)=={'Page total likes', 'Post Hour', 'Post Weekday'}

```

#### Part-2 [16 points]: We will now use neural networks to model a regression problem. 

#### 1. Create a new dataframe `df_sub4` that contains only these feature columns in `df` : 
- `Page total likes`, `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`, `Paid`,  `Total Interactions`

#### 2. Perform one-hot encoding on these features below in the dataframe `df_sub4` to create a new dataframe `df_OHE`
- `Type`, `Category`, `Post Month`, `Post Weekday`, `Post Hour`

#### 3. Perform normalization only on `Page total likes` and `Total Interactions` column in `df_OHE` using MinMaxScaler or minmax_scale 
- note that both these functions will output the same result

#### 4. Using `df_OHE` perform train_test_split operation to build training (`X_train`, `y_train`) and testing data (`X_test`, `y_test`).
- Use test_size=0.3, random_state=42 as the parameters for train_test_split() function. 
- Feature columns (X): Everything except `Total Interactions`
- CLass label column (y): `Total Interactions`

#### 5. Build a 3-layer neural network with `relu` as the 2 hidden layers' activation function and `sigmoid` as the final layer's activation function. 
- Use the same `adam` optimizer and `mean_squared_error` as the loss function. 
- `epochs`=150, `batch_size`=10
- since the class label is continuous, we use `metrics.MeanSquaredError()` as the metric (`from keras import metrics`) to measure the performance of our classifier captured as variable `mse` 

```python
#Please run this cell to install keras
#!pip install keras
#!pip install tensorflow
```

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION


from sklearn.preprocessing import MinMaxScaler
from sklearn.preprocessing import minmax_scale
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score


df_sub4=df.loc[:, ['Page total likes', 'Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour', 'Paid', 'Total Interactions']]
#print(set(df_sub4.columns))
#print(df_sub4)

colnames=['Type', 'Category', 'Post Month', 'Post Weekday', 'Post Hour']
df_OHE = pd.get_dummies(df_sub4, columns=colnames, drop_first=False)
print(len(df_OHE.columns))
#print(df_OHE)
print(len(df_OHE))

'''
meanval=df_OHE['Page total likes'].mean()
stdval=df_OHE['Page total likes'].std()
df_OHE['Page total likes']=(df_OHE['Page total likes']-meanval)/stdval

meanval=df_OHE['Total Interactions'].mean()
stdval=df_OHE['Total Interactions'].std()
df_OHE['Total Interactions']=(df_OHE['Total Interactions']-meanval)/stdval
'''

#df_OHE[['Page total likes','Total Interactions']] = MinMaxScaler().fit_transform(df[['Page total likes','Total Interactions']])

df_OHE['Page total likes'] = minmax_scale(df_OHE['Page total likes'])
df_OHE['Total Interactions'] = minmax_scale(df_OHE['Total Interactions'])

print('Sum')
print(math.ceil(df_OHE['Page total likes'].sum()))
print(math.ceil(df_OHE['Total Interactions'].sum()))
print(math.ceil(df_OHE['Page total likes'].mean()))
print(math.ceil(df_OHE['Total Interactions'].mean()))

#print(df_OHE['Total Interactions'])

X=df_OHE.drop(columns=['Total Interactions'])
y=df_OHE['Total Interactions']

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.3, random_state=42)
print(len(X_train))
print(len(y_test))

from keras.models import Sequential
from keras.layers import Dense
from keras import metrics

model = Sequential()

model.add(Dense(12, input_dim=50, activation='relu'))
model.add(Dense(8, activation='relu'))
model.add(Dense(1, activation='sigmoid'))

model.compile(loss='mean_squared_error', optimizer='adam', metrics=[metrics.MeanSquaredError()])

model.fit(X_train, y_train, epochs=150, batch_size=10)

_, mse = model.evaluate(X_test, y_test)

print("Neural Net 3-layer accuracy: "+str(mse))

### END SOLUTION
```

<!-- grade_id: q2-test1, points: 8, grade: true, solution: false -->
```python
##Cell-5 -- Do not modify this cell
assert (math.ceil(df_OHE['Page total likes'].sum())) == 357
assert (math.ceil(df_OHE['Total Interactions'].sum())) == 17
assert (math.ceil(df_OHE['Page total likes'].mean())) == 1
assert (math.ceil(df_OHE['Total Interactions'].mean())) == 1
```

<!-- grade_id: q2-test2, points: 8, grade: true, solution: false -->
```python
##Cell-6 -- Do not modify this cell
assert len(X_train)==346
assert len(y_test)==149
assert mse<0.005
assert len(model.weights)==6
```