#### Total points: 30 

#### This will be graded towards Research Translation Exercise 

#### Please do not create nor modify any cells. Please write your code ONLY where it is required and submit your python notebook via LMS. 

```python
#Please run this cell before answering the following questions 
import pandas as pd
import numpy as np

df = pd.read_csv('https://raw.githubusercontent.com/lmanikon/lmanikon.github.io/master/teaching/datasets/RReviews.tsv', sep='\t')
df.head()
```

#### Q1 [10 points]: Data preprocessing 

<ul>
    <li> Split the data into two parts -- variable `X` that contains only features and variable `y` that contains the class label only. </li>
    <li> You will notice that `X` will be a Pandas Series by now. Please use this to convert each value in X into lowercase. Once in lowercase, please remove all the non-alphanumeric characters in each review. </li>
    <li> Create a new variable `cc` (can be a list) that includes the character count of each review. Please note that index should match with the index of reviews in `X`.</li>
    <li> Create a new variable `wc` (can be a list) that includes the word count of each review. Please note that index should match with the index of reviews in `X`.</li>
</ul>


<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
#Please do not create or delete any cells
import re

cc=[]
wc=[]
### BEGIN SOLUTION 

###First split the data into X -- feature set and y -- class label question


X = df['Review']
y = df['Liked']

X = X.apply(lambda x:re.sub(r'\W+', ' ', x.lower()))

### Extract features now 

# - Number of characters
# - Number of words
cc = X.apply(lambda x:len(x))
wc = X.apply(lambda x:len(x.split(' ')))


### END SOLUTION
```

<!-- grade_id: 1, points: 10, grade: true, solution: false -->
```python
assert type(X)==type(y)
assert len(X)==1000
assert set(y)=={0,1}
assert X[0]=='wow loved this place '
assert cc[0]== 21
assert cc[999]== 132
assert wc[0]== 5
assert wc[49]== 20
assert round(np.mean(cc),2)==57.57
assert round(np.mean(wc),2)==12.12
```

#### Q2 [10 points]: Feature Generation

<ul>
    <li>Extract TFIDF vectors using the reviews in `X` and assign them to `newX`. Please note that all the modifications on reviews in the previous question should be reflected.</li>
    <li>Use Principal Component Analysis to extract the top-300 PCs. Assign these top-300 PCs to variable `newX_pca`.</li>
    <li>Standardize both `newX` and `newX_pca` and save them as `newX_std` and `newX_pca_std` respectively using StandardScaler function.</li>
</ul>


<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
#Please do not create or delete any cells
from sklearn.feature_extraction.text import TfidfVectorizer
#Standardization 
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA

### BEGIN SOLUTION

newX = TfidfVectorizer().fit_transform(X)
newX_std = StandardScaler().fit_transform(newX.todense())

pca = PCA(n_components=300)
newX_pca = pca.fit_transform(newX.todense())
newX_pca_std = StandardScaler().fit_transform(newX_pca)

#print(len(newX.todense()))
#print(len(newX.todense()[0]))
#print(round(np.mean(newX_std),2))

#print(len(newX_pca))
#print(len(newX_pca[0]))
#print(round(np.mean(newX_pca_std),2))
### END SOLUTION
```

<!-- grade_id: 2, points: 10, grade: true, solution: false -->
```python
assert len(newX.todense())==1000
assert len(newX.todense()[0])==1
assert round(np.mean(newX_std),2)==0
assert len(newX_pca)==1000
assert len(newX_pca[0])==300
assert round(np.mean(newX_pca_std),2)==0
```

#### Q3 [10 points]: Building models using original set of features

<ul>
    <li>Use standardized data of original set of features ('newX_std') and class labels (`y`) to first perform train_test_split using `test_size`=0.3, `random_state`=10 to obtain X_train, X_test, y_train, y_test corresponding to training data of features, testing data of features, class labels for data points in training data, class labels for data points in testing data respectively.</li>
    <li>Using `liblinear` as the solver, build a logistic regression model using training data and compute its accuracy using testing data. Assign accuracy value to variable `acc_lg`. </li>
    <li>Using max_depth=4, random_state=10, build a decision tree model using training data and compute its accuracy using testing data. Assign accuracy value to variable `acc_dt`. </li>
</ul>


<!-- grade_id: q3-sol, points: 0, grade: false, solution: true -->
```python
#Please do not create or delete any cells
#Build logistic regression and decisiontree classifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score

### BEGIN SOLUTION 


X_train, X_test, y_train, y_test = train_test_split(newX_std, y, test_size=0.3, random_state=10)

lg = LogisticRegression(solver='liblinear')
lg.fit(X_train, y_train)
y_pred_lg = lg.predict(X_test)
acc_lg = (accuracy_score(y_test, y_pred_lg))

dt = DecisionTreeClassifier(max_depth=4, random_state=10)
dt.fit(X_train, y_train)
y_pred_dt = dt.predict(X_test)
acc_dt = (accuracy_score(y_test, y_pred_lg))

#print(round(acc_lg, 2))
#print(round(acc_dt, 2))

### END SOLUTION
```

<!-- grade_id: 3, points: 10, grade: true, solution: false -->
```python
assert round(acc_lg, 2)==0.77
assert round(acc_dt, 2)<=0.77
```

#### Q4 [10 points]: Building models using top-300 principal component features

<ul>
    <li>Use standardized data of modified features as the top-300 principal components (`newX_pca_std`) and class labels (`y`) to first perform train_test_split using `test_size`=0.3, `random_state`=10 to obtain X_train, X_test, y_train, y_test corresponding to training data of features, testing data of features, class labels for data points in training data, class labels for data points in testing data respectively.</li>
    <li>Using `liblinear` as the solver, build a logistic regression model using training data and compute its accuracy using testing data. Assign accuracy value to variable `acc_pca_lg`. </li>
    <li>Using max_depth=4, random_state=10, build a decision tree model using training data and compute its accuracy using testing data. Assign accuracy value to variable `acc_pca_dt`. </li>
</ul>


<!-- grade_id: q4-sol, points: 0, grade: false, solution: true -->
```python
#Please do not create or delete any cells

#Build logistic regression and decisiontree classifier
from sklearn.tree import DecisionTreeClassifier
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split
from sklearn.metrics import accuracy_score

### BEGIN SOLUTION

X_train, X_test, y_train, y_test = train_test_split(newX_pca_std, y, test_size=0.3, random_state=10)

lg = LogisticRegression(solver='liblinear')
lg.fit(X_train, y_train)
y_pred_lg = lg.predict(X_test)
acc_pca_lg = (accuracy_score(y_test, y_pred_lg))

dt = DecisionTreeClassifier(max_depth=4, random_state=10)
dt.fit(X_train, y_train)
y_pred_dt = dt.predict(X_test)
acc_pca_dt = (accuracy_score(y_test, y_pred_lg))

print(round(acc_pca_lg, 2))
print(round(acc_pca_dt, 2))

### END SOLUTION 
```

<!-- grade_id: 4, points: 10, grade: true, solution: false -->
```python
assert round(acc_pca_lg, 2)<=0.80
assert round(acc_pca_dt, 2)<=0.80
```